Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Vibhu Janardhanan"
COLLABORATORS = ""
```

---

# CSE 30 Fall 2021 - Homework 5


### Instructions

Please disregard the YOUR NAME and COLLABORATORS above.  They are put there atomatically by the grading tool.
You can find instructions on how to work on a homework on Canvas.  Here is a short summary: 

### Submitting your work

To submit your work: 

* First, click on "Runtime > Restart and run all", and check that you get no errors.  This enables you to catch any error you might have introduced, and not noticed, due to your running cells out of order. 
* Second, download the notebook in .ipynb format (File > Download .ipynb) and upload the .ipynb file to [this form](https://docs.google.com/forms/d/e/1FAIpQLSdZ5gdp-2RB6wdevdSYohKMA4nbLwYV8mrHJFMRnbJ_GdTS-Q/viewform?usp=sf_link). 

You can submit multiple times; the last submission before the deadline is the one that counts.

### Homework format

For each question in this notebook, there is: 

* A text description of the problem. 
* One or more places where you have to insert your solution.  You need to complete every place marked: 

    `# YOUR CODE HERE`
    
    and you should not modify any other place. 
* One or more test cells.  Each cell is worth some number of points, marked at the top.  You should not modify these tests cells.  The tests pass if no error is printed out: when there is a statement that says, for instance: 

    `assert x == 2`
    
    then the test passes if `x` has value 2, and fails otherwise.  You can insert a `print(x)` (for this case!) somewhere if you want to debug your work; it is up to you.  
    
### Notes:

* Your code will be tested both according to the tests you can see (the `assert` statements you can see), _and_ additional tests.  This prevents you from hard-coding the answer to the particular questions posed.  Your code should solve the _general_ intended case, not hard-code the particular answer for the values used in the tests. 

* **Please do not delete or add cells!** The test is autograded, and if you modify the test by adding or deleting cells, even if you re-add cells you delete, you may not receive credit. 

* **Please do not import modules that are not part of the [standard library](https://docs.python.org/3/library/index.html).** You do not need any, and they will likely not available in the grading environment, leading your code to fail. 

* **If you are inactive too long, your notebook might get disconnected from the back-end.** Your work is never lost, but you have to re-run all the cells before you continue. 

* You can write out print statements in your code, to help you test/debug it. But remember: the code is graded on the basis of what it outputs or returns, not on the basis of what it prints.

* **TAs and tutors have access to this notebook,** so if you let them know you need their help, they can look at your work and give you advice. 

### Grading

Each cell where there are tests is worth a certain number of points.  You get the points allocated to a cell only if you pass _all_ the tests in the cell. 

The tests in a cell include both the tests you can see, and other, similar, tests that are used for grading only.  Therefore, you cannot hard-code the solutions: you really have to solve the essence of the problem, to receive the points in a cell. 

### Code of Conduct

* Work on the test yourself, alone. 
* You can search documentation on the web, on sites such as the Python documentation sites, Stackoverflow, and similar, and you can use the results. 
* You cannot share your work with others or solicit their help.


We have seen in the chapter on classes how to implement complex numbers: 

```python
import math

class Complex(object):

    def __init__(self, r, i):
        self.r = r # Real part
        self.i = i # Imaginary part

    def __add__(self, other):
        return Complex(self.r + other.r, self.i + other.i)

    def __sub__(self, other):
        return Complex(self.r - other.r, self.i - other.i)

    def __mul__(self, other):
        return Complex((self.r * other.r - self.i * other.i), 
                       (self.r * other.i + self.i * other.r))
    
    @property
    def modulus_square(self):
        return self.r * self.r + self.i * self.i
    
    @property
    def modulus(self):
        return math.sqrt(self.modulus_square)

    def inverse(self):
        m = self.modulus_square # to cache it
        return Complex(self.r / m, - self.i / m)
    
    def __truediv__(self, other):
        return self * other.inverse()

    def __repr__(self):
        """This defines how to print a complex number."""
        if self.i < 0:
            return "{}-{}i".format(self.r, -self.i)
        return "{}+{}i".format(self.r, self.i)

    def __eq__(self, other):
        """We even define equality"""
        return self.r == other.r and self.i == other.i
```

There are several ideas above: 

* To implement the mathematical operations `+`, `-`, `*`, `/`, between complex numbers, we implement the methods `__add__`, `__sub__`, `__mul__`, `__truediv__`.  You can find more information in the documentation for the [Python data model](https://docs.python.org/3/reference/datamodel.html). 
* Similarly, to define equality we define `__eq__`, and to define `<` we define `__lt__`. 

We will now do something similar to define fractions. 

## Problem 1: Implement Fractions

We want to define a class `Fraction` to represent a fraction, with integers as numerator and denominator.  
Similarly to the `Complex` class above, you need to implement the methods necessary to define `+`, `-`, `*`, `/` among fractions, as well as equality. 

You will represent fractions in _normal form_, such that:
* numerator and denumerator which do not have common factors (common divisors), except for 1 (of course), 
* the denominator is positive.

For example, when you create a fraction via: 

    r = Fraction(8, 6)

and then ask for the denominator, 

    r.numerator 

the result will be 4, and 

    r.denominator

will be 3. 

To remove the common factors from a fraction $m/n$, simply compute the greatest common divisor $d = gcd(m, n)$ via Euclid's algorithm (see the chapter on recursion), and reduce the fraction to $(m/d) / (n/d)$.

We advise you to reduce a fraction into normal form directly into the constructor of the class (the `__init__` method).  

Here is the code for `Fraction`; we leave the interesting bits for you to do.


<!-- grade_id: cell-922b1abc0b559aa7, points: 0, grade: false, solution: true -->
```python
# Definition of Fraction class

# Here is the gcd function, as it may well be useful to you.
def gcd(m, n):
    # This is the "without loss of generality" part.
    m, n = (m, n) if m > n else (n, m)
    m, n = abs(m), abs(n)
    return m if n == 0 else gcd(m % n, n)

class Fraction(object):

    def __init__(self, numerator, denominator):
        assert isinstance(numerator, int)
        assert isinstance(denominator, int)
        assert denominator != 0
        # YOUR CODE HERE
        divisor = gcd(numerator,denominator)
        if denominator < 0:
          numerator *= -1
          denominator *= -1
        self.numerator = numerator//divisor
        self.denominator = denominator//divisor



    def __repr__(self):
        """Pretty print a fraction."""
        return "{}/{}".format(self.numerator, self.denominator)


    ## Here, implement the methods for +, -, *, /, =, and <.
    ## Done quite at leisure, with spaces and all, this can be done in about
    ## 25 lines of code.
    # YOUR CODE HERE
    @staticmethod
    def commonDenom(Frac1,Frac2):
        Frac1Num = Frac1.numerator * Frac2.denominator
        Frac2Num = Frac2.numerator * Frac1.denominator
        d = Frac1.denominator * Frac2.denominator
        return(Frac1Num,Frac2Num,d)


    def __add__(self, other):
        vals = Fraction.commonDenom(self,other)
        numerator,d = vals[0] + vals[1],vals[2]
        return Fraction(numerator,d)
    def __eq__(self, other):
        return (self.numerator == other.numerator and self.denominator == other.denominator)
    def __gt__(self, other):
        vals = Fraction.commonDenom(self, other)
        return vals[0] > vals[1]
    def __lt__(self, other):
        vals = Fraction.commonDenom(self, other)
        return vals[0] < vals[1]
    def __sub__(self, other):
        vals = Fraction.commonDenom(self, other)
        numerator, d = vals[0] - vals[1], vals[2]
        return Fraction(numerator, d)
    def __mul__(self, other):
        return Fraction(self.numerator * other.numerator,self.denominator * other.denominator)
    def __truediv__(self, other):
        return Fraction(self.numerator * other.denominator, self.denominator * other.numerator)

```

```python
## Here is an example.  Feel free also to use this cell to test your code. 
for i in range(-1000,1000):
  for j in range(-1000,1000):
    if j != 0:
      f = Fraction(i,j)
      print(i,j)
      print(f)
      assert f.numerator/f.denominator == i/j
      break



```

Output:
```
-1000 -1000
1/1
-999 -1000
999/1000
-998 -1000
499/500
-997 -1000
997/1000
-996 -1000
249/250
-995 -1000
199/200
-994 -1000
497/500
-993 -1000
993/1000
-992 -1000
124/125
-991 -1000
991/1000
-990 -1000
99/100
-989 -1000
989/1000
-988 -1000
247/250
-987 -1000
987/1000
-986 -1000
493/500
-985 -1000
197/200
-984 -1000
123/125
-983 -1000
983/1000
-982 -1000
491/500
-981 -1000
981/1000
-980 -1000
49/50
-979 -1000
979/1000
-978 -1000
489/500
-977 -1000
977/1000
-976 -1000
122/125
-975 -1000
39/40
-974 -1000
487/500
-973 -1000
973/1000
-972 -1000
243/250
-971 -1000
971/1000
-970 -1000
97/100
-969 -1000
969/1000
-968 -1000
121/125
-967 -1000
967/1000
-966 -1000
483/500
-965 -1000
193/200
-964 -1000
241/250
-963 -1000
963/1000
-962 -1000
481/500
-961 -1000
961/1000
-960 -1000
24/25
-959 -1000
959/1000
-958 -1000
479/500
-957 -1000
957/1000
-956 -1000
239/250
-955 -1000
191/200
-954 -1000
477/500
-953 -1000
953/1000
-952 -1000
119/125
-951 -1000
951/1000
-950 -1000
19/20
-949 -1000
949/1000
-948
```

Here are some tests. 

<!-- grade_id: cell-742e310cea52dcda, points: 5, grade: true, solution: false -->
```python
## 5 points: creating a fraction. 

## First, let us check that you correctly put the fraction into normal form, 
## without common factor between numerator and denominator, and with a 
## positive denominator. 

f = Fraction(8, 6)
assert f.numerator == 4 and f.denominator == 3

f = Fraction(-8, 6)
assert f.numerator == -4 and f.denominator == 3

f = Fraction(8, -6)
assert f.numerator == -4 and f.denominator == 3

f = Fraction(-8, -6)
assert f.numerator == 4 and f.denominator == 3

f = Fraction(0, 10)
assert f.numerator == 0 and f.denominator == 1
```

<!-- grade_id: cell-73c9861b753d3d2d, points: 5, grade: true, solution: false -->
```python
### 5 points: hidden tests for fraction creation. 

```

<!-- grade_id: cell-dff03362ea54c7f6, points: 5, grade: true, solution: false -->
```python
## 5 points: tests for fraction operations. 

f = Fraction(8, 6) + Fraction(25, 20)
assert f.numerator == 31 and f.denominator == 12
assert f == Fraction(31, 12)
assert f == Fraction(62, 24)

assert Fraction(6, 4) + Fraction(-8, 6) == Fraction(6, 4) - Fraction(8, 6)
assert not (Fraction(6, 4) + Fraction(-8, 6) == Fraction(6, 5) - Fraction(8, 6))
```

<!-- grade_id: cell-d791a21b7f5c7f13, points: 5, grade: true, solution: false -->
```python
## 5 points: Hidden tests for fraction operations. 

```

<!-- grade_id: cell-cfb833074aea2ce1, points: 5, grade: true, solution: false -->
```python
## 5 points: more tests for fractions operations. 

assert Fraction(3, 2) * Fraction(2, 3) == Fraction(1, 1)
assert Fraction(3, 2) / Fraction(2, 3) == Fraction(9, 4)
assert Fraction(3, 2) / Fraction(6, 4) == Fraction(1, 1)
assert Fraction(32, 16) == Fraction(2, 1)
assert not Fraction(33, 16) == Fraction(4, 2)

```

<!-- grade_id: cell-b925a190bb4feb95, points: 5, grade: true, solution: false -->
```python
## 5 points: More hidden tests for fraction operations. 

```

<!-- grade_id: cell-abbef6182691ce6e, points: 5, grade: true, solution: false -->
```python
## 5 points: tests for fraction comparison. 

assert Fraction(5, 7) < Fraction(5, 6)
assert Fraction(-3, 2) < Fraction(0, 3)

```

<!-- grade_id: cell-4ab2023b7daf237c, points: 5, grade: true, solution: false -->
```python
## 5 points: hidden tests for fraction comparisons. 

```

<!-- grade_id: cell-db87b55b4217f4fb, points: 10, grade: true, solution: false -->
```python
## 10 points: Let's check you leave things unchanged. 

a = Fraction(7, 8)
b = Fraction(-4, 5)
a + b
a / b
a < b
a * b
assert a == Fraction(7, 8)
assert b == Fraction(-4, 5)
```

<!-- grade_id: cell-e379d9a0ac4667a2, points: 10, grade: true, solution: false -->
```python
## 10 points: And finally, some random tests. 

import random
for _ in range(1000):
    a = Fraction(random.randint(-200, 200), random.randint(1, 100))
    b = Fraction(random.randint(-200, 200), random.randint(1, 100))
    c = Fraction(random.randint(-200, 200), random.randint(1, 100))
    assert Fraction(-1, 1000) < (a - b) * (a - b) 
    assert (a - b) * (a + b) == a * a - b * b
    z = Fraction(0, 1) # Zero, as a fraction. 
    if not ((a == z) or (b == z) or (c == z)):
        assert (a / b) * b == (a / c) * c
        assert (a / b) * (a / c) == (a * a) / (b * c)
        assert (a / b) / (b / c) == (a * c) / (b * b)
        assert (a * a * b * c) / (a * c) == a * b
```

## Question 2: An Int class

To define the value 7, you can write `Fraction(14, 2)` or `Fraction(7, 1)` (it's the same), but this is a bit inconvenient.  Write a subclass `Int` of `Fraction`, so that `Int(7)` generates a fraction with value 7. 

<!-- grade_id: cell-ef9cc4d8b5545325, points: 0, grade: false, solution: true -->
```python
class Int(Fraction):

    # YOUR CODE HERE
    def __init__(self,num):
      super().__init__(num,1)

```

```python
## You can test your solution here. 

import random
for _ in range(10000):
    a,c=random.randint(-200,200),random.randint(-200,200)
    b,d=0,0
    while b*d == 0:
        b,d = random.randint(-200, 200),random.randint(-200, 200)
    F1=Fraction(a,b)
    F2=Fraction(c,d)
    assert F1.numerator/F1.denominator==a/b
    assert F2.numerator/F2.denominator==c/d
    assert (F1-F2)*(F1+F2)==F1*F1-F2*F2
    z = Fraction(0, 1)
    if not (F1==z or F2==z):
        assert (F1 / F2) * (F1 / F2) == (F1 * F1) / (F2 * F2)
        assert (F1 / F2) / (F2 / F1) == (F1 * F1) / (F2 * F2)
        assert (F1 * F1 * F2 * F2) / (F1 * F2) == F1 * F2
    if not F2 == z:
        assert (F1 / F2) * F2 == F1
        assert (F2 / F2) * F2 == F2
        assert (F2 / F2) == Int(1)
        assert Int(10) / (Int(10)*F2) == Fraction(F2.denominator, F2.numerator)
        assert (F1 + F2)*(F1 - F2) + F2*F2 == F1*F1
        assert z == F1 - F1 + F2 - F2 + F1*z + F2*z
        assert F2 == F2*F2/F2*Int(10)/Int(1000)*Int(100) - Int(10000) + Fraction(100000,10)
        assert F1 * Fraction(a,d) == Fraction(a*a,b*d)
    F1=Int(a)
    F2=Int(c)
    assert F1.numerator/F1.denominator==a
    assert F2.numerator/F2.denominator==c
    assert (F1 - F2)*(F1+F2) == F1*F1-F2*F2
    z=Fraction(0, 1)
    if not (F1==z or F2==z):
        assert (F1 / F2)*(F1 / F2)==F1*F1/(F2*F2)
        assert (F1 / F2)/(F2 / F1)==F1*F1/(F2*F2)
        assert (F1*F1*F2*F2)/(F1*F2)==F1*F2
    if not F2==z:
        assert (F1/F2)*F2==F1
        assert (F2/F2)*F2==F2
        assert (F2/F2)==Int(1)
        assert Int(10)/(Int(10)*F2)==Fraction(F2.denominator,F2.numerator)
        assert (F1 + F2)*(F1-F2)+F2*F2==F1*F1
        assert z == F1 - F1 + F2 - F2 + F1*z + F2*z
        assert F2 == F2*F2/F2*Int(10)/Int(1000)*Int(100) - Int(10000) + Fraction(100000,10)
        assert F1 * Fraction(a,d) == Fraction(a*a,d)
        assert Int(-10) * Fraction(1,10) * F1 == F2 - F1 - F2 + z - z + F1 - F1
        assert Int(100) * Fraction(1,100) * F1 == z + F1 + z - F2 + F2 - F1 + F1
assert Fraction(8, 8) + Fraction(25, 20) == Fraction(45,20)
assert Fraction(8, 8) - Fraction(25, 20) == Fraction(5,-20)
assert Fraction(8, 8) * Fraction(25, 20) == Fraction(25*20,400)
assert Fraction(8, 8) / Fraction(25, 20) == Fraction(20,25)
assert Fraction(8, 8) / Fraction(20, 20) == Fraction(1,1)

assert Fraction(8, 1) == Int(8)
assert Fraction(-8, 1) == Int(-8)
assert Fraction(0, 100) == Int(0)
assert Fraction(0, 100) == Fraction(0,10)

assert not Fraction(7, 7) < Fraction(6, 6)
assert not Fraction(0, -2) < Fraction(0, 3)
assert Fraction(-7, 7) < Fraction(6, 6)
assert Fraction(7, -7) < Fraction(3, 3)
assert Fraction(1, 2) < Fraction(1,1)
```

And now for some tests. 

<!-- grade_id: cell-f5c77bb7d139b319, points: 5, grade: true, solution: false -->
```python
## 5 points: tests for int class. 

assert Int(3) / Int(2) == Fraction(3, 2)
assert Int(3) * Int(4) / (Int(5) + Int(2)) == Fraction(12, 7)
assert Int(3) * Int(4) / (Int(5) + Int(1)) == Fraction(2, 1)
```